# Q4

This question will walk you through the process of how a coding assignment in JupyterHub will work.

### A

There's one question below--a small code snippet. Don't worry about what it does; instead, focus on the **double-quotes**.

See how they're right next to each other? Put something between those quotes. Anything at all. The autograder for this question will literally look to see if there's so much as a single blank space between the quotes, and will call the answer correct. `" "` would be correct; so would `"this is a lot more than just a space"`. Make something up! But whatever you do, *don't delete the quotes*.

 1. Copy the *whole* line that looks like: `instring = ""`
 2. Paste it on top of the line that looks like: `raise NotImplementedError`
 3. Put something inside the double-quotes! Anything!
 
**Step 2 is crucial, otherwise your code will fail!**

<!-- grade_id: q4a, points: 0, grade: false, solution: true -->
```python
instring = ""

### BEGIN SOLUTION

### END SOLUTION

# IGNORE EVERYTHING BELOW THIS LINE!!! THANKS!!!

def things_and_stuff(some_string):
    return "Here's what you wrote: {}".format(some_string)
out = things_and_stuff(instring)
```

<!-- grade_id: q4a_tests, points: 20, grade: true, solution: false -->
```python
assert type(out) == str
assert len(out) > 0
```

### B

Here's the bread and butter of the course: you'll be doing your assignments in Jupyter notebooks.

With that done, all that's left is to submit your code. Go back to the "Assignments" tab. Notice that, in the "Downloaded assignments" tab, there's a little arrow next to "A0". Click it, and a drop-down will appear with a handy "Validate" button on the far right. This button is awesome--it runs all the autograder tests before you even submit! So you automatically know what [autograded] questions work, and which don't. Of course, there can (and will!) still be manually-graded questions, but it should give you a good idea how you're doing.

Go ahead and click validate. Provided you followed the directions, you should see this popup:

![successkid](http://cs.uga.edu/~squinn/courses/su16/csci1360e/assets/successkid.png)

All tests passed, well done! All that's left is to send the assignment back to me so I can officially record your grade. Hit the blue "Submit" button on the right. Once it goes through, notice the timestamp that shows up:

![timestamp](http://cs.uga.edu/~squinn/courses/su16/csci1360e/assets/timestamp.png)

This is the server's officially-recorded time of completion for your assignment. You can submit as many times as you want! Only the last submission will ever be graded.

**This step is crucial. If you don't submit, I will have no way of knowing that you completed your assignment.**

### Did everything work with submitting A0?

<!-- grade_id: q4b, points: 5, grade: true, solution: true -->
